Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
# Your name here
NAME = "Ian McGregor"
```

```python
## Adding a check for version
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."
```

---

# Assignment 3: Twenty Plus One Game

This assignment asks you to implement an intelligent agent to play the twenty plus one game against one or more adversaries.

Twenty plus one is a "drinking game":  before you solve this assignment the course staff would like to remind you to drink responsibly, only if you are over twenty one, and always in moderation. Remember, as a computer scientist you will depend on your brain cells to make a living and Alcohol kills brain cells.

## Twenty Plus One.

Read about the game here: https://en.wikipedia.org/wiki/21_(drinking_game)

#### Game Rules (Please Read Carefully)

- The game is played among 2 or more players. 
- We always assume player number 0 is the intelligent agent. The other players are player numbered $1$, $2$, ..., $n-1$.
- The game starts with player 0 and the initial __count__ is at 0.
  - The initial direction of play is $+1$. I.e, the the game moves from player 0, followed by 1 and so on until $n-1$. After $n-1$, it wraps around to $0$.
  - Often, the direction of play can be reversed to $-1$. In this case, the game moves from player $i$ to $i-1$, until it reaches $0$ at which point it wraps around to player $n-1$.
- Each player has one of three moves:
  - Increment the __count__ by one. In this case, the move passes to the next player in the current direction.
  - Increment the __count__ by two. In this case, the direction of play is reversed. I.e, if we were moving from player i to i+1, we will now move from player i to i-1.
  - Increment the __count__ by three. In this case, the direction of play remains unchanged but the next player in the current direction skips a turn.
- If a player $j$ causes the count to be greater than or equal to $21$, the game stops and player $j$ loses. They are supposed to take a drink. For instance, if count is 19 and player plays a "2" or "3", they lose since count increased to 21 or more.
- We will assume the drink is "CSCI 3202 Professor" brand koolaid (completely non-alcoholic but quite intoxicating nevertheless :-)
  - Players imbibing too much of this koolaid risk going to graduate school in order to  become CS professors.

### Example Game

Player 0 starts with __count__ = 0 and __dir__ = +1.
There are three players numbered  player 0, player 1 and player 2.

- Player 0 counts one.
  - New state is __count__ = 1, __dir__ = +1, player = 1.
- Player 1 counts three.
  - New state is __count__ = 4, __dir__ = +1, player = 0. Note: Player 2 skips a turn. The play wraps around to Player 0.
- Player 0 counts three.
  - New state is __count__ = 7, __dir__ = +1, player = 2. Note: Player 1 skips a turn and the play moves to Player 2.
- Player 2 counts two.
  - New State is __count__ = 9, __dir__ = -1, player = 1. Note: play reverses direction.
- Player 1 counts three.
  - New State is __count__ = 12,  __dir__ = -1, player = 2. Note: play wraps back to player 2, since player 0 skips a turn.
- Player 2 counts two.
  - New State is __count__ = 14, __dir__ = +1, player = 0.
- Player 0 counts three.
  - New State is __count__ = 17, __dir__ = +1, player = 2. 
- Player 2 counts three.
  - New State is __count__ = 20,  __dir__ = +1, player = 1. 
- Player 1 counts two.
  - New State is __count__ = 22. But since play crosses 21, player 1 loses the game and is forced to drink the Professor's CSCI 3202 Koolaid.


## Assignment At a Glance

Problem 1: Program minmax search for the adversarial version of the game.

 - 1 (A) : You need to think through the 2 player version and prove that the player who moves first can always force a win no matter what the other player does. Manually graded.
 - 1 (B) : You are provided an implementation of a class `TwentyOneGameBoard` that implements a game board. Please read through it carefully and use the functions defined in it.
   - Implement the function: `multiplayer_minmax_search`: see section 5.2.2 of the book.
 - 1 (C) : not graded -- try out your game and see how it works.
 - 1 (D) : graded -- describe how to design utility function for the cooperative case when multiple players "gang up" against one of the players.
 
 
Problem 2: Program expectimax search when there are random players as well as adversarial players.

  - 2 (A) Complete the implementation of the `utility` function.
  - 2 (B) Complete the implementation of `expectimax_search` function.
  - 2 (C) Not graded: try to play against the intelligent agent you programmed in 2 (B) and check that it performs optimally.

## Problem 1 : Adversarial Twenty One (25 points)

We assume that 21 is played between $n \geq 2$ players each of whom derives an individual utility of -1 if they end up losing the game and +1 if they end up winning the game.


### Part A (7 points)

Assume 2 players play the game and player 0 gets to play first. 

Show that player 0 can always force player 1 to lose, no matter what actions player 1 take. To do so, briefly describe the strategy they would use. Your answer __should not exceed five sentences__ and there is a very simple strategy that can guarantee player 0 the win.

<!-- grade_id: cell-a1039ea0dab85e40, points: 7, grade: true, solution: true -->

- Player 0 counts 3
  - New state is __count__ = 3, __dir__ = +1, player = 0. Note: Player 1 skips a turn. The play wraps around to Player 0.

- Player 0 counts 3
  - New state is __count__ = 6, __dir__ = +1, player = 0. Note: Player 1 skips a turn. The play wraps around to Player 0.

- Player 0 counts 3
  - New state is __count__ = 9, __dir__ = +1, player = 0. Note: Player 1 skips a turn. The play wraps around to Player 0.

- Player 0 counts 3
  - New state is __count__ = 12, __dir__ = +1, player = 0. Note: Player 1 skips a turn. The play wraps around to Player 0.

- Player 0 counts 3
  - New state is __count__ = 15, __dir__ = +1, player = 0. Note: Player 1 skips a turn. The play wraps around to Player 0.

- Player 0 counts 3
  - New state is __count__ = 18, __dir__ = +1, player = 0. Note: Player 1 skips a turn. The play wraps around to Player 0.

- Player 0 counts 2
  - New state is __count__ = 20, __dir__ = -1, player = 1. Note: Player 1 finally plays.

- Player 1 counts (any)
  - New state is __count__ = 21+, __dir__ = +/-1, player = 0. Note: No matter what player 1 does here the count will be 21+, they will be forced to drink that Koolaid!!! 

  -Essentially what is happening here is player 0 is skipping player 1 by incrementing by 3 every play until the count reaches 18, where player 0 will count 2 and the count will be incremented to 20, where player 1 has no choice but to lose, as any of the actions permissable will lead to their doom.

## Part B (13 Points)

We now ask you to complete a multiplayer extension of the minmax search algorithm for this game. To begin with we will examine an implementation of the game's state, actions, and rules. This will be used to program the search.

### (i) Implementation of the Game Board

We have alreaddy provided an implementation of the game board that simulates the state of the game in terms of the current count, the current direction and the player. It implements the possible actions available to each player and the results of these actions. 

First carefully read the methods available.
  - `get_current_player`: returns the player whose turn it is currently.
  - `get_current_number`: returns the count of the game.
  - `is_terminal`: checks if a game state represents a game that has ended.
  - `utility`: returns a list of numbers that represent the utility of a state for each player. This function can only be called on a game state that is terminal
  - `actions`: returns a list of actions available to the player.
  - `next_state`: for a given action returns a new state as a new instance of a  `TwentyOneGameBoard` object.

Please read through the code below carefully. You will have to use it: do not attempt to modify any of the code below.

```python
import copy
import math
class TwentyOneGameBoard:
    # The class has the following fields:
    #   - self.num_players : how many players are in the game must be >= 2.
    #   - self.game_state: A tuple with three things (count, player, direction)
    #     - count ~~> a number signifying the current count.
    #     - player ~~~> whose turn is it (between 0 and num_players-1)
    #     - direction ~~> +1 or -1

    def __init__(self, num_players, start_num = 0, start_player=0, dir = 1):
        """Constructor for the class"""
        assert(num_players >= 2)
        self.num_players = num_players
        self.game_state = (start_num, start_player, dir) # Specify number we are at, current player and game direction

    def get_current_player(self):
        """return the id of the current player"""
        (_,  player, _) = self.game_state
        return player 

    def get_current_number(self):
        """return the current count for the game."""
        (num, _, _)= self.game_state
        return num 
    
    def is_terminal(self):
        """Is this game state a terminal state?"""
        (num, _, _) = self.game_state
        return num >= 21

    def utility(self):
        """Calculate the utility of the game for each player.
            Returns a list of same size as the number of players.
        """
        (num, player, dir) = self.game_state
        if self.is_terminal():
            # Game is in a terminal state
            # Utility for the player who lost is -1
            #         for everyone else who did not lose is +1
            return  [ -1 if (j==player) else 1 
                     for j in range(self.num_players) ]
        else: 
            assert False, "Utility function is not directly defined for non-terminal states"
            return [0]*self.num_players

        
    def actions(self):
        """Return the set of actions available"""
        if self.is_terminal():
            return [] # Game over - no actions
        else:
            return [1, 2, 3] # Game is not over: actions available are 1, 2 or 3
    
    def pretty_print(self):
        """ Pretty Print the game state."""
        (num, player, dir)  = self.game_state
        print('Number: %d , Player to play: %d, direction: %d' % (num, player, dir))
        if self.is_terminal():
            print('Player %d has LOST! :-( ' % player)
        return 
    
    def next_state(self, action):
        """Calculate next state for a given action"""
        (num, player, dir)  = self.game_state
        assert not self.is_terminal()
        next_num = num + action
        next_player = player 
        next_dir = dir 
        if num + action >= 21:
            # Game will end as a result of this action.
            # Prepare a new game object accordingly
            new_obj = copy.copy(self) # make a copy of myself
            # Just keep the same player and increment to the next number
            new_obj.game_state  = (next_num, next_player, next_dir)
            assert new_obj.is_terminal()
            return new_obj
        else:
            if action == 1:
                # Move play to next player in the current game direction
                next_player = player + dir  
                next_dir = dir # Next direction remains the same
            elif action == 2:
                # reverse game direction and move game to next player in the reversed direction
                next_player  = player - dir 
                next_dir = - dir 
            else: 
                # skip one player
                assert action == 3
                next_player = player + 2 * dir
                next_dir = dir # same direction
            # implement wrap around
            if next_player  < 0:
                next_player = self.num_players + next_player
            elif next_player >= self.num_players:
                next_player = next_player - self.num_players
            assert (0 <= next_player < self.num_players)
            # use the copying trick since this method may be called from a derived class
            # create a fresh instance of the game board and update its game state.
            new_obj = copy.copy(self)
            new_obj.game_state  = (next_num, next_player, next_dir)
            return new_obj
```

```python
## Illustrate the example at the beginning.
# Let us test to see that we get the expected behavior
state = TwentyOneGameBoard(3) # A game with three players
state.pretty_print() 
action_list = [1,3,3,2,3,2,3,3,2] # list of actions
for a in action_list:
    print('- Player %d plays move %d' % 
          (state.get_current_player(), a))
    assert a in state.actions(), "Illegal action for the current state"
    # compute the next state
    state = state.next_state(a)
    print('\t', end='')
    state.pretty_print()
```

Output:
```
Number: 0 , Player to play: 0, direction: 1
- Player 0 plays move 1
	Number: 1 , Player to play: 1, direction: 1
- Player 1 plays move 3
	Number: 4 , Player to play: 0, direction: 1
- Player 0 plays move 3
	Number: 7 , Player to play: 2, direction: 1
- Player 2 plays move 2
	Number: 9 , Player to play: 1, direction: -1
- Player 1 plays move 3
	Number: 12 , Player to play: 2, direction: -1
- Player 2 plays move 2
	Number: 14 , Player to play: 0, direction: 1
- Player 0 plays move 3
	Number: 17 , Player to play: 2, direction: 1
- Player 2 plays move 3
	Number: 20 , Player to play: 1, direction: 1
- Player 1 plays move 2
	Number: 22 , Player to play: 1, direction: 1
Player 1 has LOST! :-(
```

#### (ii) Implement Search

Your goal is to complete the implementation of the function `multiplayer_minmax_search` on a game which implements the multiplayer version of the search algorithm.

<!-- grade_id: cell-b518382251584341, points: 0, grade: false, solution: true -->
```python
def multiplayer_minmax_search(state):
    cur_player = state.get_current_player()
    if state.is_terminal():
        # State is dead-end state, directly return the utlity
        return ( state.utility(), None ) # there is no best action, so return None
    else: 
        # get a list of actions enabled.
        actions_enabled = state.actions()
        best_utility = [-float('inf')]*state.num_players
        best_action = None
        for a in actions_enabled:
            # 1. compute next state obtained through action a on current state
            next_state = state.next_state(a)
            #     2. Recursively calculate the utility of this next state through a 
            #        recursive call back to multiplayer_minmax_search
            next_utility, next_action = multiplayer_minmax_search(next_state)
            #     3. compare the utility of the current player from the list in step 2 
            #        to the utility of current player in the  best_utility.
            if(next_utility[cur_player] > best_utility[cur_player]):
                # 3.1 if the utility for current player improves upon the best so far, 
                # Update best so far
                best_utility = next_utility
                best_action = next_action
        return (best_utility, best_action)
```

<!-- grade_id: cell-b22cc69a518f5233, points: 6, grade: true, solution: false -->
```python
print('----- Test 1 ----- ')
b1 = TwentyOneGameBoard(2, 20, 0, 1) # Force the state to count = 20, player = 0 and direction = 1
b1.pretty_print()
(util, act) = multiplayer_minmax_search(b1)
print('Your code return utility:', util)
assert util == [-1, 1], "Test 1 failed"

print('----- Test 2 ------ ')
b2 = TwentyOneGameBoard(2, 18, 0, 1) # Force the state to count = 18, player = 0, direction = 1
b2.pretty_print()
(util2, act2) = multiplayer_minmax_search(b2)
print('Your code return utility:', util2)
assert ( util2 == [1, -1])

print('----- Test 3 ------ ')
b3 = TwentyOneGameBoard(2, 17, 0, 1) # Force the state to count = 17, player = 0, direction = 1
# Note that player 1 will win from this state, no matter what player 0 does.
b3.pretty_print()
(util3, act3) = multiplayer_minmax_search(b3)
print('Your code return utility:', util3)
assert ( util3 == [-1, 1])


print('All tests passed: 6 points')
```

Output:
```
----- Test 1 ----- 
Number: 20 , Player to play: 0, direction: 1
Your code return utility: [-1, 1]
----- Test 2 ------ 
Number: 18 , Player to play: 0, direction: 1
Your code return utility: [1, -1]
----- Test 3 ------ 
Number: 17 , Player to play: 0, direction: 1
Your code return utility: [-1, 1]
All tests passed: 6 points
```

<!-- grade_id: cell-4b34fb5a4f3030ed, points: 4, grade: true, solution: false -->
```python
print('----- Test 4 ----- ')
b0 = TwentyOneGameBoard(2) 
b0.pretty_print()
(util, act) = multiplayer_minmax_search(b0)
print('Your code return utility:', util)
assert util == [1, -1], "Test 4 failed"
print('Tests passed: 4 point')
```

Output:
```
----- Test 4 ----- 
Number: 0 , Player to play: 0, direction: 1
Your code return utility: [1, -1]
Tests passed: 4 point
```

<!-- grade_id: cell-bdc1f6dc53457936, points: 3, grade: true, solution: false -->
```python
print('----- Test 5 ----- ')
b0 = TwentyOneGameBoard(3) 
b0.pretty_print()
(util, act) = multiplayer_minmax_search(b0)
print(util)
print('Your code return utility:', util)
assert util[0] == 1, "Test 5 failed"
print('Tests passed: 3 point')
```

Output:
```
----- Test 5 ----- 
Number: 0 , Player to play: 0, direction: 1
[1, 1, -1]
Your code return utility: [1, 1, -1]
Tests passed: 3 point
```

### Part (C): Not graded
Try to play the game with the intelligent agent.

```python
# Interactively play the game with the users
def play_basic_twenty_one_game(num_players):
    board = TwentyOneGameBoard(num_players)
    while not board.is_terminal():
        print('--------------------------------')
        board.pretty_print()
        current_player = board.get_current_player()
        if current_player == 0:
            best_utility, player_move  = multiplayer_minmax_search(board)
            print('\t Computer move : +%d' % player_move)
            print('Minimax-utility of that move for various players is: %s ' % str(best_utility))
        else:
            player_move = 0
            while player_move < 1 or player_move > 3:
                player_move = int(input('\t Enter player # %d move [1-3]:' % current_player))
        board = board.next_state(player_move)
    assert board.is_terminal()
    print('Player %d got a drink!! ' %board.get_current_player() )
    if (board.get_current_player() == 0):
        print('The humans beat me!')
    else: 
        print('Intelligent agent wins: I exist to get others drunk on the CSCI 3202 koolaid.')


```

```python
play_basic_twenty_one_game(3)
```

Output:
```
--------------------------------
Number: 0 , Player to play: 0, direction: 1
<ipython-input-12-e6e7e2905779> in <module>
----> 1 play_basic_twenty_one_game(3)

<ipython-input-11-6083dce5e8c8> in play_basic_twenty_one_game(num_players)
      8         if current_player == 0:
      9             best_utility, player_move  = multiplayer_minmax_search(board)
---> 10             print('\t Computer move : +%d' % player_move)
     11             print('Minimax-utility of that move for various players is: %s ' % str(best_utility))
     12         else:

TypeError: %d format: a number is required, not NoneType
```

### Part (D) 5 Points

Thus far, the objective of the game was for each player to avoid getting a drink. The utilities modeled this fact. Suppose we now wish to model a different version where the players "gang up" against the intelligent agent (player 0), where the players wish to cooperate to ensure that the intelligent agent gets a drink while the intelligent agent would like to avoid the drink.

How would you design utilities to ensure that this happens? Your answer should specify what the utility of each type of player is (intelligent agent, and players 1,.., n-1) for the following outcomes:
  - Intelligent agent gets to drink.
  - Some player other than intelligent agent gets to drink.

<!-- grade_id: cell-97fb7f37505f99ad, points: 5, grade: true, solution: true -->
In order to ensure the Intelligent Agent loses we would need to set the utilities to favor the other players. I.e. IA would need negative utilities always and the other players positive. The converse is true to ensure another player drinks other than the IA. 

## Part 2 : With Randomized Agents (25 points)

In this part, we will include agents that are so "drunk" (on koolaid, of course) that they are unable to think strategically. Whenever such a player gets a turn they will randomly select from one of the moves available to them with equal probability. The game therefore has three types of players:
  - Intelligent Agent: Player # 0 who gets to start.
  - Adversarial player: Player who wishes to avoid drinking and gets a kick out of getting others take a drink.
  - Random player: Player who plays randomly. Derives zero utility from all of the game outcomes.
  
The utilities are as follows:
  - If a random player gets a drink everyone has utility of 0. 
  - If an adversarial/intellgent agent gets a drink then utilities are as follows:
    - Intelligent/Adversarial Player: +1 if they do not get a drink, -1 if they are forced to drink at the end of the game.
    - Random Player: 0 utility no matter whether they get to drink or not.
  
### Part 2(A): Define Utility Function (3 points)

We are provided a class that extends the data structure from Part 1. 

- The class `TwentyOneBoardWithRandom` adds the extra feature of having random players on top of the class `TwentyOneGameBoard` defined previously in Part 1.

- Extra methods include:
  - `is_random_player` : given `player_id` return True if it is ID of a random player and False otherwise.
 
- Complete the implementation of the `utility` function for the class below.

<!-- grade_id: cell-1a1f669600c1835c, points: 0, grade: false, solution: true -->
```python
from random import choice as random_choice
class TwentyOneBoardWithRandom(TwentyOneGameBoard):
    # This class extends from TwentyOneGameBoardd
    # Additional Fields:
    #  random_players: A list of player ids corresponding to random players.
    
    def __init__(self, num_players, random_players):
        super().__init__(num_players)
        # the computer cannot be a random player
        assert(all( 1 <= i < num_players for i in random_players))
        self.random_players = random_players
    
    def is_random_player(self, player_id):
        return player_id in self.random_players

    def utility(self):
        (num, loser, dir) = self.game_state
        if num >= 21:
            if self.is_random_player(loser):
                # everyone gets a utility of 0
                return [0]*self.num_players
            else:
                # Whoever lost gets a utility of -1, all non-random players get a utility of +1
                # random players get a utility of 0
                utility = [0]*self.num_players
                for i in range(0, self.num_players):
                    if(self.is_random_player(i)):
                        utility[i] = 0
                    elif(not self.is_random_player(i)):
                        utility[i] = 1
                utility[loser] = -1
                return utility
        else: 
            assert False, 'Asking for utility in a non-terminal state: not well-defined in this version of minimax search'
            return [0]*self.num_players
```

<!-- grade_id: cell-9117dc375d13f9e3, points: 3, grade: true, solution: false -->
```python
### BEGIN TESTS

print('-- Test 1 -- ')
b1 = TwentyOneBoardWithRandom(3,[2])
b1.game_state= (22, 2, 1)
b1.pretty_print()
assert (b1.utility() == [0,0,0]), "Print utility must be [0,0,0]"

print('--- Test 2 ---')
b1.game_state = (21, 0, 1)
b1.pretty_print()
assert(b1.utility() == [-1, 1, 0])

print('--- Test 3 ---')
b1.game_state = (21, 1, 1)
b1.pretty_print()
assert(b1.utility() == [1, -1, 0])

print('All tests passed: 3 points')
### END TESTS
```

Output:
```
-- Test 1 -- 
Number: 22 , Player to play: 2, direction: 1
Player 2 has LOST! :-( 
--- Test 2 ---
Number: 21 , Player to play: 0, direction: 1
Player 0 has LOST! :-( 
--- Test 3 ---
Number: 21 , Player to play: 1, direction: 1
Player 1 has LOST! :-( 
All tests passed: 3 points
```

### Part 2 (B) Program Expect-Max Search (22 points)

Complete the extension of the expectimax algorithm for dealing with the random actions of the player.

<!-- grade_id: cell-953e61120b7d3b1d, points: 0, grade: false, solution: true -->
```python
def expectimax_search(state): 
    # return the best action and corresponding utility
    # for play starting from this state
    # from the point of view of the current player.
    (_, cur_player, _) = state.game_state
    assert(0 <= cur_player < state.num_players)

# TODO: 
        #   For each action a enabled in this state
        #     1.1 compute next state on this action.
        #     1.2 Recursively call expectimax search 
        #     1.3. If cur_player is a RANDOM player
        #             Sum up the utilities from each action
        #         Else
        #             If it is better than best so far, replace the best so far.
        #   If the player is random, divide the best utility by number of actions to 
        #            compute average utility
        #   Return tuple of
        #          best utility (average utility for random player) and 
        #          corresponding action (None for random player)


    if state.is_terminal():
        # State is dead-end state
        return (state.utility(), None)
    else: 
        # get a list of actions enabled.
        actions_enabled = state.actions()
        best_action = -5
        best_utility = None # initialize best utility to None
        # Create sum_utilities list, to store and sum the utilities for each action when player is random
        sum_utilities = [0]*state.num_players
        #  1. For each action a enabled in this state
        for a in actions_enabled:
            # 1.1 compute next state on this action.
            next_state = state.next_state(a)
            # 1.2 Recursively call expectimax search 
            next_utility, next_action = expectimax_search(next_state)
            # 1.3. If cur_player is a RANDOM player
            if(state.is_random_player(cur_player)):
                # Sum up the utilities from each action
                for i in range(len(sum_utilities)):
                    sum_utilities[i] += next_utility[i]    
            else:
                # If it is better than best so far, replace the best so far.
                # Also, must check if best_utility has no value thus far, will cause issues.
                if((best_utility == None) or (next_utility[cur_player] > best_utility[cur_player])):
                    best_utility = next_utility
                    best_action = a
        # If the player is random
        if(state.is_random_player(cur_player)):
            # divide the best utility by number of actions to compute average utility
            for i in range(len(sum_utilities)):
                sum_utilities[i] = sum_utilities[i] / len(actions_enabled)
            best_utility = sum_utilities
            best_action = None
        return (best_utility, best_action)
```

<!-- grade_id: cell-7cea7e8e205e619d, points: 5, grade: true, solution: false -->
```python
### BEGIN TEST
print('-- Test 1 -- ')
b1 = TwentyOneBoardWithRandom(3,[2])
b1.game_state= (19, 2, 1)
b1.pretty_print()
(utils, act) = expectimax_search(b1)
print(utils)
assert (-0.33333333333 - 1E-05 <= utils[0] <= -0.33333333333+ 1E-05), "utilities returned by expectimax must be [-1/3,1/3,0]"
assert (0.33333333333 - 1E-05 <= utils[1] <= 0.33333333333+ 1E-05), "utilities returned by expectimax must be [-1/3,1/3,0]"
assert ( -1E-05 <= utils[2] <= 1E-05), "utilities returned by expectimax must be [-1/3,1/3,0]"
print('Test 1 Passed: 5 points')
### END TEST
```

Output:
```
-- Test 1 -- 
Number: 19 , Player to play: 2, direction: 1
[-0.3333333333333333, 0.3333333333333333, 0.0]
Test 1 Passed: 5 points
```

<!-- grade_id: cell-2bf80feaebd2a4aa, points: 5, grade: true, solution: false -->
```python
## BEGIN TESTS
print('-- Test 2 -- ')
b1 = TwentyOneBoardWithRandom(3,[2])
b1.game_state= (18, 2, 1)
b1.pretty_print()
(utils, act) = expectimax_search(b1)
print(utils)
assert (0.6666666 - 1E-04 <= utils[0] <= 0.6666666+ 1E-04), "utilities returned by expectimax must be [-1/3,1/3,0]"
assert (-0.6666666 - 1E-04 <= utils[1] <= 0.6666666+ 1E-04), "utilities returned by expectimax must be [-1/3,1/3,0]"
assert ( -1E-05 <= utils[2] <= 1E-05), "utilities returned by expectimax must be [-2/3,2/3,0]"
print('Test 2 Passed: 5 points')
## END TESTS
```

Output:
```
-- Test 2 -- 
Number: 18 , Player to play: 2, direction: 1
[0.6666666666666666, -0.6666666666666666, 0.0]
Test 2 Passed: 5 points
```

<!-- grade_id: cell-6042119167dd41ad, points: 7, grade: true, solution: false -->
```python
### BEGIN TESTS
print('-- Test 3 -- ')
b1 = TwentyOneBoardWithRandom(4,[2])
b1.game_state= (18, 3, 1)
b1.pretty_print()
(utils, act) = expectimax_search(b1)
print(utils)
assert utils == [1,-1,0,1]
assert act == 1
print('Test 3 Passed: 7 points')
### END TESTS
```

Output:
```
-- Test 3 -- 
Number: 18 , Player to play: 3, direction: 1
[1, -1, 0, 1]
Test 3 Passed: 7 points
```

<!-- grade_id: cell-e41e67fc70096a86, points: 5, grade: true, solution: false -->
```python
### BEGIN TESTS
print('-- Test 4 -- ')
b1 = TwentyOneBoardWithRandom(4,[2])
b1.game_state= (17, 2, 1)
b1.pretty_print()
(utils, act) = expectimax_search(b1)
print(utils)

assert( -0.3333 - 1E-04 <= utils[0] <= -0.3333 + 1E-04)
assert( 0.3333 - 1E-04 <= utils[1] <= 0.3333 + 1E-04)
assert( -1E-04 <= utils[2] <= 1E-04)
assert( 1 - 1E-04 <= utils[3] <= 1 + 1E-04)
#assert utils == [1,-1,0,1]
#assert act == 1
print('Test 5 Passed: 5 points')
### END TESTS
```

Output:
```
-- Test 4 -- 
Number: 17 , Player to play: 2, direction: 1
[-0.3333333333333333, 0.3333333333333333, 0.0, 1.0]
Test 5 Passed: 5 points
```

### Part C : Not graded 

Here is the script to play the game with a random player. Try to see if you can beat the computer in a 3 player game with 1 random player and 1 adversarial player. Post the "transcript" of the session on piazza if you managed to beat the computer. Let us know how.

```python
def play_random_twenty_one_game(num_players, random_player_ids):
    board = TwentyOneBoardWithRandom(num_players, random_player_ids)
    while not board.is_terminal():
        print('--------------------------------')
        board.pretty_print()
        current_player = board.get_current_player()
        if current_player == 0:
            best_utility, player_move = expectimax_search(board)
            print('\t Computer move : +%d' % player_move)
            print('Minimax-utility of that move for various players is: %s ' % str(best_utility))
        elif board.is_random_player(current_player):
            act = board.actions()
            player_move = random_choice(act)
            print('Random player chooses action %d' % player_move)
        else:
            player_move = 0
            while player_move < 1 or player_move > 3:
                player_move = int(input('\t Enter player # %d move [1-3]:' % current_player))
        board = board.next_state(player_move)

    assert board.is_terminal()
    print('Player %d got a drink!! ' % board.get_current_player() )
    if (board.get_current_player() == 0):
        print('The adversaries beat me $%^^&&!!. I will get them next time')
    elif board.is_random_player(board.get_current_player()):
        print('Another random player gets the drink! ')
    else: 
        print('Intelligent agent wins: I exist to get my adversaries drunk on the CSCI 3202 koolaid.')

```

```python
play_random_twenty_one_game(3, [1])
```

#### That's All Folks!